<font color="Teal" face="Comic Sans MS,arial">
  <h1 align="center"><i>Números de punto flotante y Técnicas de redondeo</i></h1>
  </font>
  <font color="Black" face="Comic Sans MS,arial">
  <h5 align="center"><i>Profesora: 	Ursula Xiomara Iturrarán Viveros</i></h5>
  <h5 align="center"><i>Ayudante: Juan Pablo Cordero Santiago</i></h5>
  <h5 align="center"><i>MACTII: Edgar Dominguez Rosas y José Antonio Borras Gutiérrez</i></h5>
  <h5 align="center"><i>Materia: Análisis Numérico</i></h5>
  </font>

```python
from math import pi
from PF import truncar,rdnd
```

# **Sistemas de punto flotante** 

## Definición de un sistema de punto flotante.

**Se consideran los elementos de un sistema de punto flotante a los números que tienen la siguiente forma:**

$$x = \pm (d_{1}.d_{2}d_{3}d_{4},...,d_{t}) \times \beta^{e}$$
donde

1. $t=$ precisión del sistema de punto flotante.
2. $\beta =$ base del sistema de punto flotante.
3. $e=$ Exponente.

además se debe considerar que:

$$0 \quad < d_{1} < \quad \beta$$
$$0 \quad \leq d_{i} < \quad \beta, \quad \text{para} \quad i=2,3,4,...,t$$
$$e \in [L,U], \quad \text{donde} \quad L <  0 < U$$

Se denota como $\mathbb{F}= fl(\beta, t,L,U)$ se conoce como notación estándar.

# **¿Cómo representa los números internamente la computadora?**

**La computadora tiene memoria limitada por lo tanto, no puede representar una infinidad de números solo una cantidad limitada.**

Los números de la computadora son números de un sistema de punto
flotante, ¿Qué pasa si un número tiene expansion periódica infinita
como lo puede representar la computadora?

Pensemos que la precisión $t = 2$ (también se usa $p = 2$) y pensemos en el número $x = 16/9(= 1.7777777777777777)$, la computadora lo
aproximará a 2 dígitos, pero ¿cómo? Para esto existen dos formas, con
redondeo o truncamiento. Definamos completamente al sistema de
punto flotante hipotético al que nos referimos en este caso:

# **Redondeo y truncamiento**

$\color{white}{Este\ sistema\ lo\ denotaremos\ como:\ fl(10, 3, −3, 3),\ esto\ quiere\ decir\ que\ tiene\ base\ 10,\ que\ tiene\ 3\ d\acute{i}gitos\ de\ precisi\acute{o}n\ y\ que\ el\ exponente\ e\ se\ encuentra\ en\ el\ intervalo\ [−3, 3]}$

¿Qué elementos se encuentran en este sistema de punto flotante?

Probemos $x_{1} = 7$, sabemos que lo podemos escribir como:

$$x = 7 = 7.0 \times 10^{0}$$

Siendo la base $β = 10$ y la precisión $t = 3$ (también se denota como $p$) y el exponente, en este caso, es $e = 0 \in [−3, 3]$, el exponente de $x_{1}$ está en el interavalo $[−3, 3]$, solo tenemos un dígito que es el 7 que esta en la base $\beta = 10$, por lo tanto es un número que se expresa completamente sin ningún problema en este sistema de punto flotante dado, es decir $7.0 \in fl(10, 3, −3, 3)$ 

## Ejemplos: Representación

**Ejemplo 1**: Representación de número en un sistema

$\color{white}{¿C\acute{o}mo\ puedo\ representar\ al\ n\acute{u}mero\ x_{2} = 148,\ escribiendolo\ en\ la\
notaci\acute{o}n\ est\acute{a}ndar,\ si\ fl(10, 3, −3, 3)?}$

Tenemos que:

$$x_{2} = 1.48 \times 10^{2}$$

Como tenemos 3 dígitos de precisión, alcanza para expresar el $1.48$,
además, la base es $\beta = 10$ y $1.48$ está en base 10. Por otro lado como el exponente es $e = 2 \in [−3, 3]$, este también es un número que se puede representar exactamente en el sistema de flotante dado, es decir $x_{2} \in fl(10, 3, −3, 3)$, la computadora no necesita aproximarlo de ninguna forma, pues es exactamente representable en este sistema de punto flotante.

**Ejemplo 2**: Representar un número

$\color{white}{¿C\acute{o}mo\ puedo\ representar\ al\ n\acute{u}mero\ x_2 = 8540,\ escribiendolo\ en\ la\ notaci\acute{o}n\ est\acute{a}ndar,\ si\ fl(10, 3, −3, 3)?}$

Tenemos:

$$x_{3} = 8540.0 \times 10^{0} = 8.540 \times 10^{3} = 8.54 \times 10^{3}$$

Como el sistema tiene 3 dígitos de precisión disponibles, en este caso
el último dígito es cero, **SI** alcanza para expresar el $8.540$, además este número está en la base $\beta = 10$ y el exponente es $e = 3 \in [−3, 3]$, por lo que este número puede ser representado se forma exacta en el sistema de punto flotante $fl(10, 3, −3, 3)$.

**Ejemplo 3**: Aproximar un número

**¿Cómo puedo representar al número $x_{4} = 2543$, escribiendolo en la
notación estándar, si $fl(10, 3, −3, 3)$?**

Tenemos:

$$x_{4} = 2543.0 × 10^{0} = 2.543 × 10^{3}$$

Como el sistema tiene 3 dígitos de precisión disponibles, en este caso,
**NO** alcanza para expresar el $2.543$, aunque está en la base $\beta = 10$ y el exponente es $e = 3 \in [−3, 3]$, pero este número se tiene que truncar o redondear para poder ser *aproximado* en el sistema de punto flotante $fl(10, 3, −3, 3)$, pues solo hay 3 dígitos y tenemos 4 dígitos diferentes
de cero.

## Truncamiento

La técnica de truncamiento es la más simple, pues solo se consideran
los dígitos disponibles y lo demas se borra. Consideremos $x_{7} = 2555$, al truncar para aproximar al sistema de punto flotante $fl(10, 3, −3, 3)$, tenemos que $x_{7} = 2555 = 2.555 \times 10^{3} = 2.55 \times 10^{3}$, con truncamiento.

*Diferencia entre truncamiento y redondeo*

Cuando truncamos un número dado, siempre aproximamos con el
elemento menor o igual (perteneciente al sistema de punto flotante en
cuestión), mientras que cuando redondeamos puede ser que se
aproxime al menor o al mayor, dependiendo de su cercanía a
cualquiera de estos dos números o si se da un empate entre el mayor y
menor.

## Técnicas de redondeo y truncamiento

$\color{white}{Considerando\ el\ caso\ x_4 = 2543,\ para\ aproximarlo\ a\ alg\acute{u}n\ elemento\ de\ fl(10, 3, −3, 3),\ podemos\ redondearlo\ al\ elemento\ m\acute{a}s\ cercano\ en\ este\ sistema\ de\ punto\ flotante}$

Tenemos:

$$x_4 = 2543.0 \times 10^{0} = 2.543 \times 10^{3}$$

Se encuentra entre dos elementos $x_{l}$ y $x_{u}$ que si están en $fl(10, 3, −3, 3)$:

\begin{array}{cc}
x_{l} = 2540 < x_{4} & = & 2543 < x_{u} = 2550\\
x_{l} = 2.54 \times 10^{3} < x_{4} & = & 2.543 \times 10^{3} < x_{u} = 2.55 \times 10^{3}
\end{array}

Como tengo que redondear al más cercano entonces $fl(2453) = x_{l} = 2.54 \times 10^{3}$. Estaré entonces realizando una
aproximación que tendrá un error de aproximación.

Realizar los siguientes ejemplos de truncamiento al más cercano

* $x_{5} = 2549.99999999999999$
* $x_{6} = 2550.00000000000001$
* $x_{7} = 2555$

Consideremos $x_{5} = 2549.99999999999999$

\begin{array}{cc}
x_{l} = 2540 < x_5  & = & 2549.99999999999999 < x_{u} = 2550\\
x_{l} = 2.54 \times 10^{3} & < & 2.54999999999999999 \times 10^{3} < x_{u} = 2.55 \times 10^{3}
\end{array}

Como tengo que redondear al más cercano entonces

$fl(2549.99999999999999) = x_{u} = 2.55 \times 10^{3}$.

Consideremos $x_{6} = 2549.00000000000001$

\begin{array}{cc}
x_{l} = 2540 < x_{6} & = & 2549.00000000000001 < x_{u} = 2550\\
x_{l} = 2.54 \times 10^{3} & < & 2549.00000000000001 \times 10^{3} < x_{u} = 2.55 \times 10^{3}
\end{array}

Como tengo que redondear al más cercano entonces

$fl(2549.00000000000001) = x_{u} = 2.55 \times 10^{3}$



Consideremos $x_{7} = 2555$

\begin{array}{cc}
x_{l} = 2550 < x_{6} & = & 2555 < x_{u} = 2560\\
x_{l} = 2.55 \times 10^{3} & < & 2.555 \times 10^{3} < x_{u} = 2.56 \times 10^{3}
\end{array}

En este caso se produce un empate, pues el número a aproximar se
encuentra entre el el número inferior $x_{l}$ y superior $x_{u}$ que pertenecen sistema de punto flotante. En este caso se realiza el redondeo a pares, que significa que se redondea al elemento de punto flotante cuyo último dígito significativo sea par. En este caso:

$fl(2555.0) = 2.56 \times 10^{3} = 2560.0$

# Redondedo vs truncamiento

Para la resolver la dificultad del ejemplo anterior (ejemplo 3) se puede utilizar, como ya dijimos, el redondeo o el truncamiento. El truncamiento consiste, de forma simple, en ignorar todos aquellos digítos que no alcanzan a representarse en el sistema de punto flotante que estemos utilizando, en este caso $fl(10, 3, −3, 3)$. En el caso del redondeo, el número que ocupamos para representar el número real es aquel que se aproxime más dentro del sistema de punto flotante que hemos definido. Veamos los siguientes ejemplos continuando con el sistema $fl(10, 3, −3, 3)$.

a) Consideremos $x_{a} = 2549.99999999999999$

Se encuentra entre dos elementos $x_{l}$ y $x_{u}$ que si están en $fl(10, 3, −3, 3)$:

\begin{array}{cc}
x_{l} = 2540 < x_5  & = & 2549.99999999999999 < x_{u} = 2550\\
x_{l} = 2.54 \times 10^{3} & < & 2.54999999999999999 \times 10^{3} < x_{u} = 2.55 \times 10^{3}
\end{array}

Como tengo que redondear al más cercano entonces

$fl(2549.99999999999999) = x_{u} = 2.55 \times 10^{3}$.

b) Consideremos $x_{b} = 2555$

\begin{array}{cc}
x_{l} = 2550 < x_{6} & = & 2555 < x_{u} = 2560\\
x_{l} = 2.55 \times 10^{3} & < & 2.555 \times 10^{3} < x_{u} = 2.56 \times 10^{3}
\end{array}

En este caso se produce un empate, pues el número a aproximar se
encuentra entre el el número inferior $x_{l}$ y superior $x_{u}$ que pertenecen sistema de punto flotante. En este caso se realiza el redondeo a pares, que significa que se redondea al elemento de punto flotante cuyo último dígito significativo sea par. En este caso:

$fl(2555.0) = 2560.0 = 2.56 \times 10^{3} $

Si, en el caso de $x_{b} = 2555$, hubieramos truncado entonces la representación en punto flotante sería $2.55 \times 10^{3}$,

Tambien se han creado para este notebook dos funciones,

```
rdnd(x,n), truncar(x,n)
```
las cuales permiten redondear y truncar.

```python
#Retomando lo anterior, añade la linea para redondear a tres digitos
rdnd(2555,3)
```

```python
# Ahora truncamos,
truncar(2555,3)
```

```python
rdnd(2555,3)
```

#### Ejemplo : Uso de funciones

a) Utiliza la función $\texttt{rdnd}$, que se importo en la primer celda, para obtener la representación de $\pi$ considerando un redondeo a 4 dígitos.

b) Utiliza la función $\texttt{rdnd}$, que se importo en la primer celda, para obtener la representación de $5/3$ considerando un truncamiento a 5 dígitos.

```python
#(a)
rdnd(pi,4)
```

```python
#(b)
truncar(5/3,5)
```

#### Ejercicios:truncar y redondeo

Utiliza la función $\texttt{truncar}$, para obtener la representación de los siguintes números considerando un truncamiento a 5 dígitos.

a)$\pi$ $\quad$ b) $124.036$ $\quad$

<!-- grade_id: cell-d4a2eb99a2f0a9dc, points: 0, grade: false, solution: true -->
```python
#(a)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-91324a6fd12ef8af, points: 0, grade: false, solution: true -->
```python
#(b)
# YOUR CODE HERE
raise NotImplementedError()
```

Ejercicios

Utiliza la función $\texttt{rdnd}$, para obtener la representación de los siguintes números considerando un truncamiento a 4 dígitos.

a)$\pi$ $\quad$ b) $124.036$ $\quad$  c) $5/3$

<!-- grade_id: cell-408fa4db7fcdb0b8, points: 0, grade: false, solution: true -->
```python
#(a)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9ed201d27585d9eb, points: 0, grade: false, solution: true -->
```python
#(b)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-282d68a05fdd0ace, points: 0, grade: false, solution: true -->
```python
#(c)
# YOUR CODE HERE
raise NotImplementedError()
```

# Aritemetica de punto flotante

Consideremos la siguiente función: $f(x)=x^3-6.1x^2+3.2x+1.5$, evaluada en $f(4.71)$. Si estamos trabajando con una representación de punto flotante de los números, entonces, para cada una de las operaciones relizadas en la evaluación de dicha función se obtienen resultados aproximados, de acuerdo con los dígitos con los que contemos para representar los resultados. A continuación, se muestra cómo se vería la evaluación de la función $f$ de forma exacta, con truncamiento de 3 dígitos y con redondeo de 3 dígitos.
$$Exacta: \quad f(4.71)=104.487111-135.32301+15.072+1.5=-14.263899$$
$$Truncamiento: \quad f(4.71)=104.-134.+15.0+1.50=-13.5 \quad \qquad \qquad$$
$$Redondeo: \quad f(4.71)=105.-135.+15.1+1.50=-13.4 \quad \qquad \qquad \qquad$$

#### Ejemplo: Aritmética Flotante 

Reproduce, usando código de python, la evaluación de $f$ de usando aritmética de truncamiento de 3 dígitos.

```python
# f(x) usando truncamiento
x=4.71
eval_f=truncar( truncar(x**3,3) - truncar( 6.1*truncar(x**2,3), 3) + truncar(3.2*x,3) + 1.5 , 3 )
print(eval_f)
```

Ejemplo:

Reproduce, usando código de python, la evaluación de $f$ de usando aritmética de redondeo de 3 dígitos.

```python
x=4.71
eval_f=rdnd( rdnd(rdnd(x*x,3)*x,3) - rdnd( 6.1*rdnd(x**2,3), 3) + rdnd(3.2*x,3) + 1.5 , 3 )
print(eval_f)
```

## Error relativo
Con lo abordado en las secciones previas, es claro que, el hecho de que el cálculo de una computadora está limitado, implica que tiene **errores**. El error relativo es una forma de medir el error de un valor real '$y$' con respeto de su aproximación '$fl(y)$' y esta dada por la expresión:

$$\left|\frac{y-fl(y)}{y}\right|$$.

En muchas ocasiones el valor real '$y$' es un número abstracto que no puede ser representado por una computadora, sin embargo, para fines didácticos consideraremos que el valor real '$y$' es el valor que se obtiene al considerar toda la precicisón (dígitos) de la computadora.


Pensemos, de nuevo, en la función $f(x)=x^3-6.1x^2+3.2x+1.5$, evaluada en $f(4.71)$. Como se muestra al inicio de la sección 2.2.1, la evaluación exacta nos da $y=−14.263899$ y de la evaluación usando truncamiento se obtiene $fl(y)=-13.5$.

Entonces el error relativo es:
$$\left|\frac{-14.263899+13.5}{-14.263899}\right|\approx 0.05 \quad Truncamiento $$

```python
#---- Error relativo de truncamiento para f(4.71) ----
x=4.71
y=x**3-6.1*x**2+3.2*x+1.5 #valor 'exacto'
fly=truncar( truncar(x**3,3) - truncar( 6.1*truncar(x**2,3), 3) + truncar(3.2*x,3) + 1.5 , 3 ) #valor aproximado
er=abs( (y-fly)/y ) #error relativo
print(er)
```

#### Ejercicio: Error relativo

Encuentra el error relativo de las siguientes operaciones usando truncamiento a 4 dígitos, para ello creo u ncodigo que lo resuelva:

a)$\pi - \frac{22}{7}$ $\quad$ b) $\left( \frac{2}{9}\right) \left(\frac{9}{7} \right)$$\quad$ 

```python
#----  Error relativo de truncamiento para (a)  ----

#Inserta el valor de y para el caso propuesto
y=pi-22/7#valor 'exacto'
```

<!-- grade_id: cell-78c6cd0333fb0bd5, points: 0, grade: false, solution: true -->
```python
#Continuando con el problema, el valor flotante sera:
fly=truncar( truncar(pi,4)-truncar(22/7,4) , 4 )#valor aproximado con truncamiento

#Define el valor del error relativo, con "er" la variabel a imprimir
# YOUR CODE HERE
raise NotImplementedError()
print(er)
```

Para la parte b), definiremos de manera total una función "def" en la cual resolvera para el error relativo

<!-- grade_id: cell-7002d6a0210062be, points: 0, grade: false, solution: true -->
```python
#--- Error relativo de truncamiento para (b)  ----
def ejerciciob():
    #Escribe un codigo donde sea capaz de redondear la suma de tres numeros a dos digitos despues del punto
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
ejerciciob()
```

<!-- grade_id: cell-457f0ae2f5d6c742, points: 3, grade: true, solution: false -->
```python
assert ejerciciob() == 0.0007499999999998341, "Error en el resultado, revisa tu función"

print("El ejercicio pasa la prueba")
```

#### Recálculo de f(x) con expresión factorizada

Retomemos, una vez más, la función $f(x)=x^3-6.1x^2+3.2x+1.5$, si factorizamos algunos términos, podemos reexpressar la función como: $ f(x)=((x-6.1)x+3.2)x+1.5$. Ambas expresiones de $f$ son algebráicamente equivalentes; sin embargo, la pregunta es si obtendremos el mismo error relativo que obtuvimos anteriormente, al comienzo de la sección 2.1.3.

Recalculando, empleando la expresión factorizada tenemos:
$$Truncamiento: \quad f(4.71)=((4.71-6.1)(4.71)+3.2)(4.71)+1.5=-14.2$$
$$Redondeo: \quad f(4.71)=((4.71-6.1)(4.71)+3.2)(4.71)+1.5=-14.3$$
Así,los errores relativos seran:
$$\left|\frac{-14.263899+14.2}{-14.263899}\right|\approx 0.0045 \quad Truncamiento $$
$$\left|\frac{-14.263899+14.3}{-14.263899}\right|\approx 0.0025 \quad Redondeo $$
De este ejemplo podemos concluir que nuestros resultados han disminuido un orden de magnitud si factorizamos nuestros los valores a evaluar.

#### Ejemplo : Error elativo
Usa código de python para realizar los cálculos anteriores con redondeo de 3 dígitos.

<!-- grade_id: cell-3291eb62c78d10f2, points: 0, grade: false, solution: true -->
```python
#----  Error relativo de redondeo  ----
def ErrorRelativo():
    # YOUR CODE HERE
    raise NotImplementedError()
    return er
```

```python
ErrorRelativo()
```

<!-- grade_id: cell-41ec51e939e03233, points: 3, grade: true, solution: false -->
```python
assert ErrorRelativo() == 0.0025309349147810893, "Error en el resultado, revisa tu función"

print("El ejercicio pasa la prueba")
```

El código se ha definido como un todo, sin independencia de variables $y$ y $x$, ¿creés posible definir una función llamada "ErrorRelativo" la cual tenga dependencia de parametros como $[x,y,...]$ ? para mejorar su funcionabilidad?

```python
# Eres libre de intentarlo...
```

# Ejercicios Finales

Ejercicio a):

Explica brevemente, ¿por qué las siguientes evaluaciones lógicas, mostradas en la introducción, dan resultados erróneos?

* $(\sqrt(3))^2=3$
* $(0.1+0.1+0.1) = 0.3$
* $ 0.9999999999999999999 < 1.0 $

Ejercicio b):

El hecho que $ 0.9999999999999999999 < 1.0 $ de como resultado $\textbf{False}$, ¿se explica por una representación de punto flotante de redondeo o de truncamiento? ¿Por qué?

Ejercicio c):

Dados dos puntos $(x_0,y_0)$ y $(x_1,y_1)$ que definen una recta, tal que $y_1 \neq y_0$, se tienen dos fórmulas para encontrar la intersección-$x$:

$x=\frac{x_0 y_1 - x_1 y_0}{y_1-y_0}   \quad \quad , \quad \quad x=x_0-\frac{(x_1-x_0)y_0}{y_1-y_0} $

Usa los puntos $(x_0 , y_0 ) = (1.31, 3.24)$ y $(x_1 , y_1 ) = (1.93, 4.76)$ así como aritmética de redondeo de 3 dígitos para determinar qué método es mejor. ¿ A qué crees que se deba esta diferencia ?

<center>
<img src="https://github.com/jugernaut/Prometeo/blob/desarrollo/Figuras/gracias.png?raw=1" width="600"> 
</center>