<a href="https://colab.research.google.com/github/jugernaut/AnalisisNumerico/blob/desarrollo/01_AritmeticaPuntoFlotante/PreAritmeticaFlotanteEstudiante.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

<font color="Teal" face="Comic Sans MS,arial">
  <h1 align="center"><i>Aritmética de punto flotante</i></h1>
  </font>
  <font color="Black" face="Comic Sans MS,arial">
  <h5 align="center"><i>Profesora: 	Ursula Xiomara Iturrarán Viveros</i></h5>
    <h5 align="center"><i>Ayudante: Juan Pablo Cordero Santiago</i></h5>
  <h5 align="center"><i>Materia: Análisis Numérico</i></h5>
</font>

### Introducción
La aritmética realizada por una computadora es diferente a nuestro álgebra que utilizamos día a día, por ejemplo cuando **sumamos** $3+6=9$, o **multiplicamos** $5 \times 6 = 30$ no hay ningún problema, pero al realizar operaciones como **raíces** y **exponentes** $(\sqrt(3))^2=3$ no necesariamente obtendremos lo mismo de una computadora.

De manera computacional: cada número esta representado unicamente por un valor fijo de dígitos. Por ejemplo; solo los **números racionales** y **no pares** pueden ser expresados  exactamente. Entonces retomando el caso de $\sqrt(3)$, este es un **número no racional** y estará representado en una computadora por un número fijo de dígitos, estos logrando una buena aproximación, pero tenemos casos donde esta discrepancia se incrementa de manera indeseable.

Para ello analizemos los siguientes ejercicios de igualdad:

```python
(0.1+0.1+0.1) == 0.3
```

 <center> ¿Estas de acuerdo con la respuesta? ¿Qué crees que este pasando?</center>

 <center> Analicemos el caso ahora de $\sqrt(3))^2=3$  </center>

```python
( (3)**(0.5) )**(2) == 3
```

## Número en una computadora (64-Bits)
Una computadora funciona en base a dos dígitos **0** o **1**. Consideremos el siguiente ejemplo de 64 bits:

$$ \underset{Signo}{\underbrace{0}} \quad \underset{Parte Entera}{\underbrace{10000000011}} \quad \underset{Parte Fraccionaria}{\underbrace{1011100100010000000000000000000000000000000000000000}},$$


el primer dígito indica el signo y se denota como $s$, los siguientes $11$ dígitos se denotan como $c$ y representan un exponente y los otros $54$ son la parte fraccionaria $f$:

$$(-1)^s2^{c-1023}(1+f)$$

el $1$ que se suma a la parte fraccionaria tiene que ver con el concepto de **bit implícito**, que veremos más adelante.

Desarrollando el ejemplo pasado, obtenemos que: $s$ es cero indicando que el número es positivo, para $c$ cada dígito representa el factor multiplicando a la potencia del número $2$ elevado a la posición donde se encuentra:

$$c= 1 \cdot 2^{10}+0 \cdot 2^{9}+\cdot \cdot \cdot +0 \cdot 2^{2}+1 \cdot 2^{1}+1 \cdot 2^{0}=1024+2+1=1027$$

$f$ representa la parte fraccionaria del número, tambien denominada mantisa y está expresada como:

$$f=1\cdot\left(\frac{1}{2}\right)^1+0\cdot\left(\frac{1}{2}\right)^2+1\cdot\left(\frac{1}{2}\right)^3+1\cdot\left(\frac{1}{2}\right)^4+1\cdot\left(\frac{1}{2}\right)^5+\cdot\cdot\cdot1\cdot\left(\frac{1}{2}\right)^8 +\cdot\cdot\cdot1\cdot\left(\frac{1}{2}\right)^{12}+\cdot\cdot\cdot$$

donde de nuevo la posición del dígito indica la potencia de la fracción. Finalmente obtenemos que:

$$ \quad(-1)^s2^{c-1023}(1+f)=(-1)^02^{1027-1023}\left(1+\frac{1}{2}+\frac{1}{8}+\frac{1}{16}+\frac{1}{32}+\frac{1}{256}+\frac{1}{4096}\right)$$

$$(-1)^s2^{c-1023}(1+f)= 16 \cdot (1.722900391)$$

$$(-1)^s2^{c-1023}(1+f)= 27.56640625$$

### Limitaciones

Calculemos los dos número más cercanos a nuestro ejemplo, para el más pequeño cambiamos nuestro último $1$ a $0$ y todos los números a continuación a $1$, obteniendo:

$$0 \quad 10000000011 \quad 1011100100001111111111111111111111111111111111111111$$

y para el valor más grande, solo cambiamos el último $0$ a $1$:

$$0 \quad 10000000011 \quad 1011100100010000000000000000000000000000000000000001$$

Esto implica que nuestro número calculado no solo representa ese valor, si no la mitad de números reales entre $27.56640625$ y sus dos números más cercanos.

Marco En pocas palabras esto significa que su pudiéramos ubicar estos 3 valores en la recta numérica, podríamos ver que hay mucho espacio entre ellos, y eso se traduce en que **no sería posible representar un valor que se encuentre en estos huecos**.

## Notación Decimal

Veamos como se construye un valor en notación **decimal** ($\beta=10$);

$$\pm0.d_1 d_2 \cdot \cdot \cdot d_k \times 10^n, \quad 1< d_1 < 9, \quad 0 < d_i < 9$$

para cada $i=2,...,k.$ 

Cualquier número real y positivo en el rango de la computadora puede ser normalizado para ser expresado como:

$$\pm0.d_1 d_2 \cdot \cdot \cdot d_k d_{k+1} d_{k+2} \cdot \cdot \cdot \times 10^n$$

## Notación de punto flotante (truncamiento y redondeo)

Se denota como $fl(y)$ y se obtiene delimitando el tamaño de la **mantisa**, una de ellas puede ser simplemente **truncando** en el término $d_{k+1}$:

$$ fl(y) = 0.d_1 d_2 \cdot  \cdot  \cdot \times 10^n$$

otra manera es **redondeando**, si el dígito $k+1$ es menor a cinco, cortamos, obteniendo el mismo resultado, pero si este es cinco o mayor, se agrega el valor de uno al $k$-ésimo término.

$$fl(y)=\pm 0.d_{1}d_{2}...\tilde{d_{k}}\times10^{n} \qquad \tilde{d_{k}}=\begin{cases}
\begin{array}{c}
d_{k} \quad si\,d_{k+1}<10/2\\
d_{k}+1 \quad si\,d_{k+1}\geq 10/2
\end{array}\end{cases}$$

###  Limitaciones en *Python*

Empleando la versión de *Python* que usa *google colab* por default, responde las siguientes preguntas y actividades:

*   ¿Qué tipo de precisión se emplea en *python*?.
*   ¿Por qué la operación de la celda inferior devuelve `False`, aunque claramente debería devolver `True`?.
*   La función `round()` se emplea para redondear un valor a un cierto número de cifras significativas. Modifica el segundo *prtint* de la celda inferior para que el resultado sea `True`.

<!-- grade_id: cell-2179c16f0d85b6f8, points: 0, grade: false, solution: true -->
```python
print(.1 + .1 + .1 == .3)

print(round(.1, 1) + round(.1, 1) + round(.1, 1) == round(.3, 1))

# YOUR CODE HERE
raise NotImplementedError()
```

### Ejercicio

Define una función llamada `SumRound(a,b,c)` que sea capaz de redondear la suma de $a+b+c$ a dos dígitos después del punto decimal.

<!-- grade_id: cell-39a456858bbf6cf4, points: 0, grade: false, solution: true -->
```python
def SumRound(a,b,c):
    #Escribe un codigo donde sea capaz de redondear la suma de tres numeros a dos digitos despues del punto
    # YOUR CODE HERE
    raise NotImplementedError()
```

Vamos a comprobar que la función previamente definida es la correcta.

```python
SumRound(0.001,2.1,-2.1)
#0.001+2.1-2.1 = 0.001
```

<!-- grade_id: cell-17b2aa31e64bc5a1, points: 3, grade: true, solution: false -->
```python
"""Verificando el resultado obtenido con la función SumRound(a,b,c)"""
assert (SumRound(0.111,10.1,0.1) == 10.31 and SumRound(0.1,0.1,0.1) == 0.3 and 
        SumRound(0.001,2.1,-2.1) == 0.0), "Error en el resultado, revisa tu función"

print("La funcion SumRound pasa la prueba")
```

## Aritmética de punto flotante (Error en cómputo científico)

Con lo abordado en los temas previos, es claro que el cálculo de una computadora está limitado, lo cual implica que tiene **errores**. Uno de estos errores es el aproximar un número en su forma de punto flotante ya sea por **redondeo** o **truncamiento**.

Es por este motivo que es importante conocer las diferentes formas de medir la magnitud de un error. Sea $fl(x)$ el número de máquina generado en el sistema $\mathcal F$ asociado al valor real $x$, tendremos 2 formas de evaluar la magnitud del error.

### Error relativo
Esta forma de medir el error de un valor real con respeto de su aproximación considera la magnitud del valor real, y esta dada por la expresión.

$$\left|\frac{y-fl(y)}{y}\right|$$

De la expresión anterior podemos ver que esta normalizada, de modo que podemos generar la diferencia porcentual entre la aproximación y el valor real.


### Error absoluto

El error absoluto únicamente considera la magnitud de la diferencia entre $y$ y su aproximación $fl(y)$, es decir.

$$\left|y-fl(y)\right|$$

El error absoluto muestra la diferencia entre la aproximación $fl(y)$ y el valor $y$. En otras palabras, existe un número real $\delta$ que depende de $y$, tal que

$$fl(y)=y(1+\delta)\quad $$ siendo $$ \,\left|\delta\right|\leq u$$

Por otra parte, debido a que los números de punto flotante no están igualmente espaciados, es necesario tratar de encontrar una cota para el error absoluto.

El error absoluto está acotado por la siguiente expresión.

$$\left|y-fl(y)\right|\leq\frac{1}{2}\beta^{-n+e}$$

## Ejemplos y ejercicios


**1.   Ejemplo:**

Calcular el error relativo y absoluto del valor $\pi=3.14159265...$ al utilizar el método de truncamiento y el de redondeo a cinco dígitos (Hint: para el valor de $y$ utiliza $\pi$ a cinco dígitos):

Por truncamiento $fl(y)=0.31415\times 10^1$, y para el redondeo tenemos un nueve en el sexto digito, por lo que;
$$ fl(y)=(0.31415+0.00001)\times 10^1=0.31416\times10^1$$

El error absoluto y relativo para el método de truncamiento es cero, mientras que para el caso del redondeo:
$$\left|y-fl(y)\right|=\left| 3.1415-3.1416 \right|=0.0001$$

$$\left|\frac{y-fl(y)}{y}\right|=\left|\frac{0.0001}{0.31415\times10^1}\right|\approx 0.00031$$

**Ejercicios:**


Calcula el el error absoluto y relativo para los siguientes casos, redondeando a 4 dígitos el resultado final:


a) $y=e, \quad f(y)=2.718$     $\qquad \qquad$   b) $y=e^{10}, \quad f(y)=1400$

c) $y=\sqrt(2), \quad f(y)=1.414$   $\,\,$ $\qquad$   d) $y=9!, \quad f(y)=\sqrt(18\pi)(9/e)^9$

Ahora de ejercicio crearemos un función que calcule el error absoluto y relativo.

<!-- grade_id: cell-07d543d1aa19ec0f, points: 0, grade: false, solution: true -->
```python
def ErrorAbsolutoyRelativo(y,f):
    import math
    #y es el valor y f su aproximacion
    #define el valor Relativo, nombre de la variable "ErrorRelativo"
    # YOUR CODE HERE
    raise NotImplementedError()
    return round(ErrorRelativo, 4)
```

La segunda parte del ejercicio de va a evaluar la función diseñada con los ejemplos propuestos, cada una vale un punto. Procura ejecutar tu código para revisar que funcione correctamente.

```python
from math import e
import math
ErrorAbsolutoyRelativo(math.factorial(9),math.sqrt(18/math.pi)*(9/e)**9)
```

<!-- grade_id: cell-c350336521d99eee, points: 4, grade: true, solution: false -->
```python
from math import e
import math

"""Verificando el resultado obtenido con la función ErrorAbsolutoyRelativo(y,f)"""
assert (ErrorAbsolutoyRelativo(e,2.718) == 0.0001 and ErrorAbsolutoyRelativo(math.sqrt(2),1.414) == 0.0002 and 
        ErrorAbsolutoyRelativo(e**(10),1400) == 0.9364 and
        ErrorAbsolutoyRelativo(math.factorial(9),math.sqrt(18/math.pi)*(9/e)**9) == 0.6846), "Error en el resultado, revisa tu función"

print("La funcion ErrorAbsolutoyRelativo pasa la prueba")
```

## Bibliografía
1. Numerical Methods, Third Edition, J. Douglas Faires
2. Notas del curso de Analisis Numerico, [Git](https://github.com/jugernaut/Numerico2021), Miguel Angel Pérez Leon.
3. Riswan Butt: Numerical Analysys Using Matlab, Jones and Bartlett.