<a href="https://colab.research.google.com/github/jugernaut/AnalisisNumerico/blob/desarrollo/02_ResolucionDeSistemasLineales/Clase_NormasVectoriales.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Normas vectoriales y matriciales

## Notebook de clase

---

## Introducción

### ¿Qué es una norma vectorial?

Sea $V$ un espacio vectorial sobre $k$. Una función $\left\Vert \cdot\right\Vert$ de vectores, se denomina norma vectorial a $\left\Vert \cdot\right\Vert$ si, para cualesquiera vectores $\vec{x},\vec{y},$ se satisfacen las siguientes propiedades.

1. $\left\Vert \vec{x}\right\Vert \geq0.$
2. $\left\Vert \vec{x}\right\Vert =0\,\Longleftrightarrow\vec{x}=\vec{0}.$
3. $\left\Vert \alpha\vec{x}\right\Vert =\left|\alpha\right|\left\Vert \vec{x}\right\Vert$, donde $a$ es una constante.
4. $\left\Vert \vec{x}+\vec{y}\right\Vert \leq\left\Vert \vec{x}\right\Vert +\left\Vert \vec{y}\right\Vert.$

### Tipos de norma

Sea $p \geq 1$, las normas de Hölder, o *normas-p* están definidas por:

$$
\left\Vert\vec{x}\right\Vert_{p}=\left[\sum_{i=1}^{n}\left|x_{i}\right|^{p}\right]^{1/p} \qquad \text{para}\,1\leq p<\infty
$$

Los casos más importantes de está norma son:

* Norma-1 ($p = 1$):

$$
\left\Vert \vec{x}\right\Vert _{1}=\sum_{i=1}^{n}\left|x_{i}\right|
$$

* Norma-2 o norma euclideana ($p = 2$)

$$
\left\Vert\vec{x}\right\Vert_{2}=\left[\sum_{i=1}^{n}\left|x_{i}\right|^{2}\right]^{1/2}
$$

* Y la norma-$\infty$ o norma del supremo ($p = \infty$)

$$
\lim_{p \rightarrow \infty} \left\Vert \vec{x}\right\Vert _{p} = \left\Vert \vec{x}\right\Vert _{\infty}= \max_{1 \leq i \leq n} \left| x_i \right|
$$

### Implementando las normas-p en Python

Sea $\vec{x} \in \mathbb{R}^{2}$ y $\vec{x}=(2,5)$, calcule las normas de Hölder con p = 1, 2 y $\infty$.

* Norma-1

$$\left\Vert \vec{x}\right\Vert _{1}=\sum_{i=1}^{2}\left|x_{i}\right|=\left|x_{1}\right|+\left|x_{2}\right|=2+5=7$$

* Norma-2 (norma euclidiana)

$$\left\Vert \vec{x}\right\Vert _{2}=\left[\sum_{i=1}^{2}\left|x_{i}\right|^{2}\right]^{1/2}=\sqrt{\left|x_{1}\right|^{2}+\left|x_{2}\right|^{2}}=\sqrt{\left|2\right|^{2}+\left|5\right|^{2}}=\sqrt{4+25}=\sqrt{29}$$

* Norma-$\infty$ (norma del supremo)

$$\left\Vert \vec{x}\right\Vert _{\infty} = \max_{1\leq i\leq 2}\left|x_{i}\right|= \max\{|2|,|5|\} = \max\{2,5\}=5$$

Con *numpy* podemos comprobar los resultados:

```python
import numpy as np

# Definimos el vector
x = np.array([2, 5])

# Se imprime el vector x
print("Vector x:", x)

# Calculamos la norma 1
print("Norma 1 del vector x:",np.linalg.norm(x, 1))

# Calculamos la norma 2
print("Norma 2 del vector x:",np.linalg.norm(x, 2))

# Calculamos la norma del supremo
print("Norma del supremo del vector x:",np.linalg.norm(x, np.inf))
```

## Normas matriciales (inducida)

El concepto de norma, no se aplica únicamente para vectores, en cualquier espacio vectorial o subespacio es muy útil definir el concepto de norma para poder establecer una forma de medir distancias.

Dada una norma vectorial, definimos su correspondiente norma matricial de una matriz $\mathbf{A}$ de tamaño $m \times n$ como sigue:

$$
||\mathbf{A}|| = \max_{\mathbf{x} \not = 0}  \frac{||A \mathbf{x}||}{||\mathbf{x}||}
$$

esta norma se llama inducida o subordinada a la norma vectorial.

### Propiedades

Una función $\left\Vert \cdot\right\Vert$ de matrices, se denomina norma matricial $\left\Vert \cdot\right\Vert$  si para cualesquiera matrices $A, B \in M_{n\times n}$ sobre $\mathbb{R}$ se satisfacen las siguientes propiedades.

1. $\left\Vert \mathbf{A}\right\Vert \geq0$ si $A \neq$ 0
2. $\left\Vert \mathbf{A}\right\Vert =0\,\Longleftrightarrow A=0$.
3. $\left\Vert \alpha \mathbf{A}\right\Vert =\left|\alpha\right|\left\Vert \mathbf{A}\right\Vert$.
4. $\left\Vert \mathbf{A}+\mathbf{B}\right\Vert \leq\left\Vert \mathbf{A}\right\Vert +\left\Vert \mathbf{B}\right\Vert$ (desigualdad triangular).
5. $\left\Vert \mathbf{AB}\right\Vert \leq\left\Vert \mathbf{A}\right\Vert \left\Vert \mathbf{B}\right\Vert$ (compatibilidad o consistencia).

Nótese que las primeras dos propiedades juntas implican que $\left\Vert \mathbf{A} \right\Vert = 0$ sí y sólo sí $\mathbf{A} = 0$

### Tipos de normas matriciales

Sea $p\geq1,\:\vec{x}\in\mathbb{R}^{n}$, la p-norma de una matriz $\mathbf{A}\in \mathbf{M}_{n\times n}$ esta definida por.

$$
\left\Vert \mathbf{A}\right\Vert _{p} = \sup_{\vec{x}\neq\vec{0}}\frac{\left\Vert A\vec{x}\right\Vert _{p}}{\left\Vert \vec{x}\right\Vert _{p}}
$$

* Norma matricial inducida por la norma-1 vectorial ($p = 1$).

$$
\left\Vert \mathbf{A}\right\Vert _{1}\ = \sup_{\vec{x}\neq\vec{0}}\frac{\left\Vert A\vec{x}\right\Vert _{1}}{\left\Vert\vec{x}\right\Vert _{1}} = \max_{1\leq j\leq n}\sum_{i=1}^{n}\left|a_{ij}\right| = \max_{1\leq j\leq n}\left\Vert \vec{a}_{j}\right\Vert _{1} \quad \text{suma de columnas}
$$

* Norma matricial inducida por la norma del supremo ($p = \infty$).

$$
\left\Vert \mathbf{A}\right\Vert _{\infty} = \sup_{\vec{x}\neq\vec{0}}\frac{\left\Vert A\vec{x}\right\Vert _{\infty}}{\left\Vert \vec{x}\right\Vert _{\infty}} = \max_{1\leq i\leq n}\sum_{j=1}^{n}\left|a_{ij}\right| = \max_{1\leq i\leq n}\left\Vert \vec{a}_{i}\right\Vert _{1} \quad \text{suma de renglones}
$$

* El caso particular de $p=2$ es conocido como norma euclidiana:
$$\left\Vert \mathbf{A}\right\Vert _{2}=\sqrt{\max{\text{eig}(\mathbf{A}^{T}\mathbf{A})}}$$

* Además de estas, existe la norma de Frobenius:

$$
\left\Vert \mathbf{A} \right\Vert_{F} = \sqrt{\sum_{i = 1}^{m} \sum_{j = 1}^{n} |a_{ij}^2|} = \sqrt{\text{tr}\,\left(\mathbf{A}^{*}\mathbf{A}\right)} = \sqrt{\sum_{i = 1}^{\min{m, n}} \sigma_{i}^2}
$$

donde $\mathbf{A^{*}}$ es la traspuesta conjugada y $\sigma_{i}$ son los valores singulares de $\mathbf{A}$

### Ejemplo 1: $\left\Vert \mathbf{A} \right\Vert_{1}$ y $\left\Vert \mathbf{A} \right\Vert_{\infty}$

Sea

$$
\mathbf{A} = 
\begin{pmatrix}
3 & -1 & 4 \\
-5 & 0 & 2 \\
1 & -2 & 6 \\
\end{pmatrix}
$$

Calcular $\left\Vert \mathbf{A} \right\Vert_{1}$ y $\left\Vert \mathbf{A} \right\Vert_{\infty}$

$$
\left\Vert \mathbf{A}\right\Vert _{1} = \max_{1\leq j\leq n} \left\Vert \vec{a}_{j}\right\Vert _{1} = \max_{1\leq j\leq 3} \left\Vert \vec{a}_{j}\right\Vert _{1}= \max\{(3+5+1=9),(1+0+2=3),(4+2+6=12)\}=12
$$

$$
\left\Vert \mathbf{A}\right\Vert _{\infty}=\max_{1\leq i\leq n} \left\Vert \vec{a}_{i}\right\Vert _{1} = \max_{1\leq i\leq 3}\left\Vert \vec{a}_{i}\right\Vert _{1}= \max\{(3+1+4=8),(5+0+2=7),(1+2+6=9)\}=9
$$

Comprobando el resultado con el módulo ```linalg``` de ```numpy```:

```python
# Definiendo la matriz A
A = np.array([[3, -1, 4],[-5, 0, 2],[1, -2, 6]])
print("Matriz: \n", A)

# Calculamos la 1-norma de A
print("Norma-1:", np.linalg.norm(A, 1))

# Calculamos la norma del supremo de A
print("Norma-inf:", np.linalg.norm(A, np.inf))
```

### Ejemplo 2: $\left\Vert \mathbf{A}\right\Vert _{2}$

Sea $$\mathbf{A}=\left(\begin{array}{cc}
2 & 5\\
1 & 3
\end{array}\right)$$ Calcular $\left\Vert \mathbf{A}\right\Vert _{2}$

* Para encontrar los eigenvalores de A necesitamos calcular $det\left(\left(\mathbf{A}^{T}\mathbf{A}\right)-\lambda I\right)$.

* Lo primero es calcular la matriz transpuesta de $\mathbf{A}$, es decir $$\mathbf{A}^{T}=\left(\begin{array}{cc}
2 & 1\\
5 & 3
\end{array}\right)$$

* Ahora es necesario encontrar el producto de $\mathbf{A}$ transpuesta por $\mathbf{A}$.$$\mathbf{A}^{T}\mathbf{A}=\left(\begin{array}{cc}
2 & 1\\
5 & 3
\end{array}\right)\left(\begin{array}{cc}
2 & 5\\
1 & 3
\end{array}\right)=\left(\begin{array}{cc}
5 & 13\\
13 & 34
\end{array}\right)$$

* Lo siguiente es calcular el producto anterior menos $lambda$ por la matriz de identidad.$$\mathbf{A}^{T}\mathbf{A}-\lambda I=\left(\begin{array}{cc}
5 & 13\\
13 & 34
\end{array}\right)-\left(\begin{array}{cc}
\lambda & 0\\
0 & \lambda
\end{array}\right)=\left(\begin{array}{cc}
5-\lambda & 13\\
13 & 34-\lambda
\end{array}\right)$$

* Finalmente se calcula el determinante de la matriz anterior 
$$\begin{eqnarray*}\det\left(\left(\mathbf{A}^{T}\mathbf{A}\right)-\lambda I\right) =	\left(5-\lambda\right)\left(34-\lambda\right)-\left(13\times13\right) \\
=	170-5\lambda-34\lambda+\lambda^{2}-169 \\
=	\lambda^{2}-39\lambda+1 \end{eqnarray*}$$

* Por lo que al encontrar las raíces de esta ecuación cuadrática, se tiene que$$\begin{cases}
\lambda_{1}=\frac{39+\sqrt{1517}}{2} & =38.97434\\
\lambda_{2}=\frac{39-\sqrt{1517}}{2} & =0.02566
\end{cases}$$

$$\therefore\left\Vert \mathbf{A}\right\Vert _{2}=\sqrt{38.97434}=6.24294338$$

```python
# Definimos la matriz A
A = np.array([[2, 5],[1, 3]])
print("Matriz:\n", A)

# Calculamos la 2-norma de A
print("Norma-2:", np.linalg.norm(A, 2))
```

### Ejercicio 1

Calcula los tres tipos de normas matriciales ($\left\Vert \mathbf{A}\right\Vert _{p}$) para los siguientes ejercicios propuestos e imprime el resultado.

**NOTA**: Para este ejercicio, **SÍ** puedes utilizar el módulo ```numpy```

a) $\mathbf{A} = \begin{pmatrix}
-2 & 15 & 18\\
0 & 3 & 10\\
0 & 0 & 1
\end{pmatrix}$

<!-- grade_id: cell-a688bdb505d1e48a, points: 0, grade: false, solution: true -->
```python
# Definición de la matriz
# YOUR CODE HERE
raise NotImplementedError()

# Norma-1
# YOUR CODE HERE
raise NotImplementedError()

# Norma-2
# YOUR CODE HERE
raise NotImplementedError()

# Norma-inf
# YOUR CODE HERE
raise NotImplementedError()
```

b) $\mathbf{B} = \begin{pmatrix}
0 & 1 & 0\\
1 & 0 & 1\\
0 & 1 & 0
\end{pmatrix}$

<!-- grade_id: cell-ef98aae59040c769, points: 0, grade: false, solution: true -->
```python
# Definición de la matriz
# YOUR CODE HERE
raise NotImplementedError()

# Norma-1
# YOUR CODE HERE
raise NotImplementedError()

# Norma-2
# YOUR CODE HERE
raise NotImplementedError()

# Norma-inf
# YOUR CODE HERE
raise NotImplementedError()
```

c) $\mathbf{C} = \begin{pmatrix}
2 &  4\\
5 &  7
\end{pmatrix}$

<!-- grade_id: cell-e5c7496302cbc0da, points: 0, grade: false, solution: true -->
```python
# Definición de la matriz
# YOUR CODE HERE
raise NotImplementedError()

# Norma-1
# YOUR CODE HERE
raise NotImplementedError()

# Norma-2
# YOUR CODE HERE
raise NotImplementedError()

# Norma-inf
# YOUR CODE HERE
raise NotImplementedError()
```

d) $\mathbf{D} = \begin{pmatrix}
\frac{1}{2} &  \frac{3}{2}\\
\frac{5}{2} &  -\frac{1}{2}
\end{pmatrix}$

<!-- grade_id: cell-c834a6173ef1634e, points: 0, grade: false, solution: true -->
```python
# Definición de la matriz
# YOUR CODE HERE
raise NotImplementedError()

# Norma-1
# YOUR CODE HERE
raise NotImplementedError()

# Norma-2
# YOUR CODE HERE
raise NotImplementedError()

# Norma-inf
# YOUR CODE HERE
raise NotImplementedError()
```

## Ejercicio Final

Para reforzar lo aprendido en esta clase, deberás generar la función ```norma_frob``` la cuál calculara la norma de Frobenius para matrices.

**NOTA**: Para este ejercicio, no se puede hacer uso de ```np.linalg.norm```

<!-- grade_id: cell-0e19714de2d5b429, points: 0, grade: false, solution: true -->
```python
import numpy as np

def norma_frob(matriz):
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return norma
```

<!-- grade_id: cell-e9e87db3b09a8857, points: 100, grade: true, solution: false -->
```python
"""Prueba 1: Norma de Frobenius para matrices"""
from nose.tools import assert_equal
from numpy.linalg import norm

# Norma de Frobenius
# Caso 1
assert_equal(norma_frob(np.array([[30, 20, 10], [40, 34, 12]])), norm(np.array([[30, 20, 10], [40, 34, 12]]), 'fro'))
print("La norma de Frobenius de\n", np.array([[30, 20, 10], [40, 34, 12]]), "\nes:", norma_frob(np.array([[30, 20, 10], [40, 34, 12]])))

# Caso 2
assert_equal(norma_frob(np.array([[-8, -9], [-1, -1], [7, 4]])), norm(np.array([[-8, -9], [-1, -1], [7, 4]]), 'fro'))
print("La norma de Frobenius de\n", np.array([[-8, -9], [-1, -1], [7, 4]]), "\nes:", norma_frob(np.array([[-8, -9], [-1, -1], [7, 4]])))

# Caso 3
assert_equal(norma_frob(np.array([[0, 5, 4], [5, 0, 3], [4, 3, 0]])), norm(np.array([[0, 5, 4], [5, 0, 3], [4, 3, 0]]), 'fro'))
print("La norma de Frobenius de\n", np.array([[0, 5, 4], [5, 0, 3], [4, 3, 0]]), "\nes:", norma_frob(np.array([[0, 5, 4], [5, 0, 3], [4, 3, 0]])))

# Caso 4
assert_equal(norma_frob(np.array([[1.00, 0.97, 0.79, 0.45], [0.97, 1.00, 0.86, 0.39], [0.79, 0.86, 1.00, 0.54], [0.45, 0.39, 0.54, 1.00]])), 
             norm(np.array([[1.00, 0.97, 0.79, 0.45], [0.97, 1.00, 0.86, 0.39], [0.79, 0.86, 1.00, 0.54], [0.45, 0.39, 0.54, 1.00]]), 'fro'))
print("La norma de Frobenius de\n", np.array([[1.00, 0.97, 0.79, 0.45], [0.97, 1.00, 0.86, 0.39], [0.79, 0.86, 1.00, 0.54], [0.45, 0.39, 0.54, 1.00]]), "\nes:", 
      norma_frob(np.array([[1.00, 0.97, 0.79, 0.45], [0.97, 1.00, 0.86, 0.39], [0.79, 0.86, 1.00, 0.54], [0.45, 0.39, 0.54, 1.00]])))
```