<a href="https://colab.research.google.com/github/jugernaut/AnalisisNumerico/blob/desarrollo/02_ResolucionDeSistemasLineales/PreClase_NormasVectoriales.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Normas vectoriales y matriciales

## Notebook pre-clase

---

## Introducción

En todo espacio vectorial es de suma importancia definir una metrica, es decir, una forma de conocer la distancia entre los diferentes elementos dentro de este espacio vectorial o la distancia de cada uno de estos al origen.

El concepto de distancia es de gran importancia para el estudio de los espacios vectoriales. 

Un ejemplo es el calculo de la distancia entre un vector $\vec{x}$ y un vector $\vec{y}$ en donde estos vectores pueden representar la fuerza de coriolis que genera la tierra y la aceleracion de una bala de 35mm. Esto puede escalar si agregamos que estamos sobre una superficie no plana.

Es por este motivo que es necesario definir una forma de medicion entre distancias en cualquier espacio vectorial. A este concepto se le conoce como **norma**.

## Normas Vectoriales

Sea $V$ un espacio vectorial sobre $k$. Una función $\left\Vert \cdot\right\Vert$ de vectores, se denomina norma vectorial a $\left\Vert \cdot\right\Vert$ si, para cualesquiera vectores $\vec{x},\vec{y},$ se satisfacen las siguientes propiedades.

1. $\left\Vert \vec{x}\right\Vert \geq0.$
2. $\left\Vert \vec{x}\right\Vert =0\,\Longleftrightarrow\vec{x}=\vec{0}.$
3. $\left\Vert \alpha\vec{x}\right\Vert =\left|\alpha\right|\left\Vert \vec{x}\right\Vert$, donde $a$ es una constante.
4. $\left\Vert \vec{x}+\vec{y}\right\Vert \leq\left\Vert \vec{x}\right\Vert +\left\Vert \vec{y}\right\Vert.$

### Tipos de normas (Hölder)

Sea $p\geq1$, las normas de *Hölder*, o *normas-p* están definidas por

$$
\left\Vert\vec{x}\right\Vert_{p}=\left[\sum_{i=1}^{n}\left|x_{i}\right|^{p}\right]^{1/p}
\qquad \text{para}\,1\leq p<\infty
$$

De ellas, las mas importantes son:

* La $\text{Norma}-1$

$$
\left\Vert \vec{x}\right\Vert _{1}=\sum_{i=1}^{n}\left|x_{i}\right|
$$

* La $\text{Norma}-2$ (norma euclidiana = espacio plano)

$$
\left\Vert\vec{x}\right\Vert_{2}=\left[\sum_{i=1}^{n}\left|x_{i}\right|^{2}\right]^{1/2}
$$

* y la $\text{Norma}-\infty$ (norma del supremo)

$$
\lim_{p \rightarrow \infty} \left\Vert \vec{x}\right\Vert _{p} = \left\Vert \vec{x}\right\Vert _{\infty}= \max_{1 \leq i \leq n} \left| x_i \right|
$$

Observacion: En la norma euclidiana podemos observar como es el calculo que normalmente utilizamos para calcular la distancia entre dos puntos en un espacio plano.

### Ejemplo

Sea $\vec{x} \in\mathbb{R}^{2}$, $\vec{x}=(2,5)$, calcule las normas de Hölder.

* Norma-1

$$\left\Vert \vec{x}\right\Vert _{1}=\sum_{i=1}^{2}\left|x_{i}\right|=\left|x_{1}\right|+\left|x_{2}\right|=2+5=7$$

* Norma-2 (norma euclidiana)

$$\left\Vert \vec{x}\right\Vert _{2}=\left[\sum_{i=1}^{2}\left|x_{i}\right|^{2}\right]^{1/2}=\sqrt{\left|x_{1}\right|^{2}+\left|x_{2}\right|^{2}}=\sqrt{\left|2\right|^{2}+\left|5\right|^{2}}=\sqrt{4+25}=\sqrt{29}$$

* Norma-$\infty$ (norma del supremo)

$$\left\Vert \vec{x}\right\Vert _{\infty} = \max_{1\leq i\leq 2}\left|x_{i}\right|= \max\{|2|,|5|\} = \max\{2,5\}=5$$

### Ejercicio 1

Usando las definiciones para los tipos de normas y basádose en el ejemplo anterior, escribe la función ```norma1``` que calcule dicha norma para vectores.

**NOTA**: Resolver el ejercicio sin usar el módulo ```numpy```

<!-- grade_id: cell-05da1ca36f6875bc, points: 0, grade: false, solution: true -->
```python
def norma1(vector):
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return norma
```

<!-- grade_id: cell-2b8dc53b0dabc206, points: 10, grade: true, solution: false -->
```python
"""Prueba 1: Función para la norma-1"""
from nose.tools import assert_equal
from numpy.linalg import norm

assert_equal(norma1([2, 4]), norm([2, 4], 1), "Error: Revisa tu función")
print("Norma-1 del vector", [2, 4], ":", norma1([2, 4]))
assert_equal(norma1([3, 0, 12, 4]), norm([3, 0, 12, 4], 1), "Error: Revisa tu función")
print("Norma-1 del vector", [3, 0, 12, 4], ":", norma1([3, 0, 12, 4]))
```

<!-- grade_id: cell-8441580043fb74d6, points: 10, grade: true, solution: false -->
```python
"""Prueba 2: Función para la norma-1"""
from nose.tools import assert_equal
from numpy.linalg import norm

assert_equal(norma1([-2, 0, 2, -4]), norm([-2, 0, 2, -4], 1), "Error: Revisa tu función")
print("Norma-1 del vector", [-2, 0, 2, -4], ":", norma1([-2, 0, 2, -4]))
assert_equal(norma1([-2.43, 23.67, 34.5, -50, 12, 32.4, -5]), norm([-2.43, 23.67, 34.5, -50, 12, 32.4, -5], 1), "Error: Revisa tu función")
print("Norma-1 del vector", [-2.43, 23.67, 34.5, -50, 12, 32.4, -5], ":", norma1([-2.43, 23.67, 34.5, -50, 12, 32.4, -5]))
```

### Ejercicio 2

Siguiendo el ejercicio 1, escribe ahora la función ```norma2``` que calcule la norma-2 para vectores.

**NOTA**: Resolver el ejercicio sin usar el módulo numpy

<!-- grade_id: cell-785ddb55557b8e5d, points: 0, grade: false, solution: true -->
```python
def norma2(vector):
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return norma
```

<!-- grade_id: cell-9fcab781df214f4a, points: 10, grade: true, solution: false -->
```python
"""Prueba 3: Función para la norma-2"""
from nose.tools import assert_equal
from numpy.linalg import norm

assert_equal(norma2([2, 4]), norm([2, 4]), "Error: Revisa tu función")
print("Norma-2 del vector", [2, 4], ":", norma2([2, 4]))
assert_equal(norma2([3, 0, 12, 4]), norm([3, 0, 12, 4]), "Error: Revisa tu función")
print("Norma-2 del vector", [3, 0, 12, 4], ":", norma2([3, 0, 12, 4]))
```

<!-- grade_id: cell-fe8cab83615339b6, points: 10, grade: true, solution: false -->
```python
"""Prueba 4: Función para la norma-2"""
from nose.tools import assert_equal
from numpy.linalg import norm

assert_equal(norma2([-2, 0, 2, -4]), norm([-2, 0, 2, -4]), "Error: Revisa tu función")
print("Norma-2 del vector", [-2, 0, 2, -4], ":", norma2([-2, 0, 2, -4]))
assert_equal(norma2([-2.43, 23.67, 34.5, -50, 12, 32.4, -5]), norm([-2.43, 23.67, 34.5, -50, 12, 32.4, -5]), "Error: Revisa tu función")
print("Norma-2 del vector", [-2.43, 23.67, 34.5, -50, 12, 32.4, -5], ":", norma2([-2.43, 23.67, 34.5, -50, 12, 32.4, -5]))
```

### Ejercicio 3

Siguiendo las definiciones dada para las normas y los ejercicios realizados con anterioridad, escribe la función ```norma_inf``` que calcule la norma del supremo para vectores.

**NOTA**: Resolver el ejercicio sin hacer uso del módulo ```numpy```

<!-- grade_id: cell-bf7ace22816786fe, points: 0, grade: false, solution: true -->
```python
def norma_inf(vector):
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return norma
```

<!-- grade_id: cell-b722b0277b30169c, points: 10, grade: true, solution: false -->
```python
"""Prueba 5: Función para la norma-inf"""
from nose.tools import assert_equal
from numpy.linalg import norm
from numpy import inf

assert_equal(norma_inf([2, 4]), norm([2, 4], inf), "Error: Revisa tu función")
print("Norma-Inf del vector", [2, 4], ":", norma_inf([2, 4]))
assert_equal(norma_inf([3, 0, 12, 4]), norm([3, 0, 12, 4], inf), "Error: Revisa tu función")
print("Norma-Inf del vector", [3, 0, 12, 4], ":", norma_inf([3, 0, 12, 4]))
```

<!-- grade_id: cell-b96fa0abb68ba6c6, points: 10, grade: true, solution: false -->
```python
"""Prueba 6: Función para la norma-inf"""
from nose.tools import assert_equal
from numpy.linalg import norm
from numpy import inf

assert_equal(norma_inf([-2, 0, 2, -4]), norm([-2, 0, 2, -4], inf), "Error: Revisa tu función")
print("Norma-Inf del vector", [-2, 0, 2, -4], ":", norma_inf([-2, 0, 2, -4]))
assert_equal(norma_inf([-2.43, 23.67, 34.5, -50, 12, 32.4, -5]), norm([-2.43, 23.67, 34.5, -50, 12, 32.4, -5], inf), "Error: Revisa tu función")
print("Norma-Inf del vector", [-2.43, 23.67, 34.5, -50, 12, 32.4, -5], ":", norma_inf([-2.43, 23.67, 34.5, -50, 12, 32.4, -5]))
```

### Ejercicio 4

Como ejercicio final, y siguiendo la definición general para la norma de Hölder, escribe la función ```norma_p```.

**NOTA**: Resolver sin utilizar el módulo ```numpy```

<!-- grade_id: cell-a6a95a62f52a9c20, points: 0, grade: false, solution: true -->
```python
def norma_p(vector, p):
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return norma
```

<!-- grade_id: cell-09b3fa682a2a04f1, points: 10, grade: true, solution: false -->
```python
"""Prueba 7: Función para la norma-p"""
from nose.tools import assert_equal
from numpy.linalg import norm

assert_equal(norma_p([2, 4], 4), norm([2, 4], 4), "Error: Revisa tu función")
print("Norma-p del vector", [2, 4], ":", norma_p([2, 4], 4))
assert_equal(norma_p([3, 0, 12, 4], 8), norm([3, 12, 4], 8), "Error: Revisa tu función")
print("Norma-p del vector", [3, 12, 4], ":", norma_p([3, 12, 4], 8))
```

<!-- grade_id: cell-7114d5cde20c969f, points: 10, grade: true, solution: false -->
```python
"""Prueba 8: Función para la norma-p"""
from nose.tools import assert_equal
from numpy.linalg import norm

assert_equal(norma_p([-2, 0, 2, -4], 3), norm([-2, 0, 2, -4], 3), "Error: Revisa tu función")
print("Norma-p del vector", [-2, 0, 2, -4], ":", norma_p([-2, 0, 2, -4], 3))
assert_equal(norma_p([-2.43, 23.67, 34.5, -50, 12, 32.4, -5], 13), norm([-2.43, 23.67, 34.5, -50, 12, 32.4, -5], 13), "Error: Revisa tu función")
print("Norma-p del vector", [-2.43, 23.67, 34.5, -50, 12, 32.4, -5], ":", norma_p([-2.43, 23.67, 34.5, -50, 12, 32.4, -5], 13))
```

### Bola (Esfera) Unitaria

Es importante notar, cómo es que la elección de la norma, modifica la forma en la que se percibe el espacio (ó subespacio) vectorial en cuestión.

Para cada norma, se dibuja a manera de ilustración la bola de unidad definida como.

$$\bar{B}\left(0,1\right):=\left\{ \vec{x}\in\mathbb{R}^{2}\mid\left\Vert \vec{x}\right\Vert _{p}\leq1\right\}$$

<center> <img src="https://github.com/jugernaut/Numerico2021/blob/desarrollo/Figuras/Normas/p1.png?raw=1" width="150"> <img src="https://github.com/jugernaut/Numerico2021/blob/desarrollo/Figuras/Normas/p2.png?raw=1" width="150">  <img src="https://github.com/jugernaut/Numerico2021/blob/desarrollo/Figuras/Normas/pinf.png?raw=1" width="150"> <img src="https://github.com/jugernaut/Numerico2021/blob/desarrollo/Figuras/Normas/pmedios.png?raw=1" width="150"> </center>

Lo que muestra la bola unitaria en realidad son todos los $\vec{x}\in R^{2}$
cuya distancia al origen es menor ó igual que 1, es decir que podemos
pensar en la bola unitaria de la siguiente manera.
$$\bar{B}\left(0,1\right)=\left\{ \vec{x}\in\mathbb{R}^{2}\mid d_{p}\left(\vec{x}-\vec{0}\right)\leq1\right\} $$

## Equivalencia de normas

Las normas son herramientas básicas para definir y analizar la
convergencia de una sucesión de vectores en espacios vectoriales. La
convergencia depende de las definición de norma que se usa. Podría
suceder que una sucesión de vectores convergiera para una norma, pero no
para otra. Pero en un espacio vectorial de dimensión finita todas las
normas son equivalentes en el siguiente sentido:

Dos normas $\left\Vert\cdot\right\Vert, \left\Vert\cdot\right\Vert'$ son equivalentes si existen $m$ y $M$ constantes reales positivas tales que para cualquier $x$ tenemos

$$
m||x||  \leq  ||x||^{'} \leq M ||x||
$$

## Distancia entre dos vectores

Sean $\vec{x}$, $\vec{y}$ $\in \mathbb{R}^{n}$ la distancia entre $\vec{x}$, $\vec{y}$ se define de la siguiente manera.

$$d\left(\vec{x},\vec{y}\right) = \left\Vert \vec{x}-\vec{y}\right\Vert$$

La importancia de una norma es que una vez definida la norma
en $\mathbb{R}^{n}$ es sencillo establecer una métrica (distancia entre vectores).

Si consideramos un espacio bidimensional y la norma-2 tenemos que:

$$
d_{2}\left(\vec{x},\vec{y}\right) = \left\Vert \vec{x}-\vec{y}\right\Vert _{2}= \left[\sum_{i=1}^{n}\left|x_{i}-y_{i}\right|^2\right]^{1/2}
$$

### Ejercicio 5

Escribe una función llamada ```distancia``` que calcule la distancia euclideana entre dos vectores cualesquiera, siguiendo la definición anterior.

**NOTA**: Resolver sin usar el módulo ```numpy```

<!-- grade_id: cell-698682ddeb68f127, points: 0, grade: false, solution: true -->
```python
def distancia(x, y):
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return dist
```

<!-- grade_id: cell-81b42969431db7a3, points: 5, grade: true, solution: false -->
```python
"""Prueba 9: Función para la distancia euclideana"""
from nose.tools import assert_equal
from numpy import array
from numpy.linalg import norm

assert_equal(distancia([6, 3], [6, 3]), norm(array([6, 3]) - array([6, 3])), "Error: Revisa tu función")
print("Distancia entre", [6, 3], "y", [6, 3], "=", distancia([6, 3], [6, 3]))
assert_equal(distancia([6, 3], [2, 2]), norm(array([6, 3]) - array([2, 2])), "Error: Revisa tu función")
print("Distancia entre", [6, 3], "y", [2, 2], "=", distancia([6, 3], [2, 2]))
assert_equal(distancia([3, 4], [6, 3]), norm(array([3, 4]) - array([6, 3])), "Error: Revisa tu función")
print("Distancia entre", [3, 4], "y", [6, 3], "=", distancia([3, 4], [6, 3]))
assert_equal(distancia([-8, 1], [5, -2]), norm(array([-8, 1]) - array([5, -2])), "Error: Revisa tu función")
print("Distancia entre", [-8, 1], "y", [5, -2], "=", distancia([-8, 1], [5, -2]))
```

<!-- grade_id: cell-9fbb32351875086b, points: 5, grade: true, solution: false -->
```python
"""Prueba 10: Función para la distancia euclideana"""
from nose.tools import assert_equal
from numpy import array
from numpy.linalg import norm

assert_equal(distancia([-10, 3, -3], [5, 4, -2]), norm(array([-10, 3, -3]) - array([5, 4, -2])), "Error: Revisa tu función")
print("Distancia entre", [-10, 3, -3], "y", [5, 4, -2], "=", distancia([-10, 3, -3], [5, 4, -2]))
assert_equal(distancia([20, -32, 43], [4, 3, 2]), norm(array([20, -32, 43]) - array([4, 3, 2])), "Error: Revisa tu función")
print("Distancia entre", [20, -32, 43], "y", [4, 3, 2], "=", distancia([20, -32, 43], [4, 3, 2]))
```

## Ejercicio Final

Para asentar los conocimientos que has adquirido para las normas vectoriales, tu trabajo será demostrar las propiedades para las 3 normas que se vieron en este documento (norma-1, norma-2 y norma-$\infty$).

1. $\left\Vert \vec{x}\right\Vert \geq0.$
2. $\left\Vert \vec{x}\right\Vert =0\,\Longleftrightarrow\vec{x}=\vec{0}.$
3. $\left\Vert \alpha\vec{x}\right\Vert =\left|\alpha\right|\left\Vert \vec{x}\right\Vert$, donde $a$ es una constante.
4. $\left\Vert \vec{x}+\vec{y}\right\Vert \leq\left\Vert \vec{x}\right\Vert +\left\Vert \vec{y}\right\Vert.$

En este ejercicio **SÍ** se puede utilizar el módulo ```numpy```.

Para hacer uso de este módulo, se debe ejecutar la siguiente línea:

```python
import numpy as np
```

y para hacer uso de la norma, la función se implementa como:

```python
np.linalg.norm(x, p)
```

donde ```x``` es el vector y ```p``` es la norma que se quiera calcular (en el caso de la norma del supremo, p = ```np.inf```)

<!-- grade_id: cell-4c95da8fd0a363fe, points: 10, grade: true, solution: true -->
```python
# 1. ||x|| >= 0
# YOUR CODE HERE
raise NotImplementedError()

# 2. ||x|| = 0 <=> x = 0
# YOUR CODE HERE
raise NotImplementedError()

# 3. ||ax|| = ||a|| ||x||, a es constante
# YOUR CODE HERE
raise NotImplementedError()

# 4. ||x + y|| <= ||x|| + ||y||
# YOUR CODE HERE
raise NotImplementedError()

```