<a href="https://colab.research.google.com/github/jugernaut/AnalisisNumerico/blob/desarrollo/01_AnalisisNumerico/Ursula/02_ResolucionDeSistemasLineales/NormasVecMat/Profesor/PostClase_NormasVectorialesYMatriciales_Profesor.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Normas Vectoriales y Matriciales

## Notebook post-clase

---

## Recapitulación

### Normas vectoriales

Sea $p \geq 1$, las normas de Hölder, o *normas-p* están definidas por:

$$
\left\Vert\vec{x}\right\Vert_{p}=\left[\sum_{i=1}^{n}\left|x_{i}\right|^{p}\right]^{1/p} \qquad \text{para}\,1\leq p<\infty
$$

Los casos más importantes de está norma son:

* Norma-1 ($p = 1$):

$$
\left\Vert \vec{x}\right\Vert _{1}=\sum_{i=1}^{n}\left|x_{i}\right|
$$

* Norma-2 o norma euclideana ($p = 2$)

$$
\left\Vert\vec{x}\right\Vert_{2}=\left[\sum_{i=1}^{n}\left|x_{i}\right|^{2}\right]^{1/2}
$$

* Y la norma-$\infty$ o norma del supremo ($p = \infty$)

$$
\lim_{p \rightarrow \infty} \left\Vert \vec{x}\right\Vert _{p} = \left\Vert \vec{x}\right\Vert _{\infty}= \max_{1 \leq i \leq n} \left| x_i \right|
$$

### Normas matriciales

Sea $p\geq1,\:\vec{x}\in\mathbb{R}^{n}$, la p-norma de una matriz $\mathbf{A}\in \mathbf{M}_{n\times n}$ esta definida por.

$$
\left\Vert \mathbf{A}\right\Vert _{p} = \sup_{\vec{x}\neq\vec{0}}\frac{\left\Vert A\vec{x}\right\Vert _{p}}{\left\Vert \vec{x}\right\Vert _{p}}
$$

* Norma matricial inducida por la norma-1 vectorial ($p = 1$).

$$
\left\Vert \mathbf{A}\right\Vert _{1}\ = \sup_{\vec{x}\neq\vec{0}}\frac{\left\Vert A\vec{x}\right\Vert _{1}}{\left\Vert\vec{x}\right\Vert _{1}} = \max_{1\leq j\leq n}\sum_{i=1}^{n}\left|a_{ij}\right| = \max_{1\leq j\leq n}\left\Vert \vec{a}_{j}\right\Vert _{1} \quad \text{suma de columnas}
$$

* Norma matricial inducida por la norma del supremo ($p = \infty$).

$$
\left\Vert \mathbf{A}\right\Vert _{\infty} = \sup_{\vec{x}\neq\vec{0}}\frac{\left\Vert A\vec{x}\right\Vert _{\infty}}{\left\Vert \vec{x}\right\Vert _{\infty}} = \max_{1\leq i\leq n}\sum_{j=1}^{n}\left|a_{ij}\right| = \max_{1\leq i\leq n}\left\Vert \vec{a}_{i}\right\Vert _{1} \quad \text{suma de renglones}
$$

* El caso particular de $p=2$ es conocido como norma euclidiana:
$$\left\Vert \mathbf{A}\right\Vert _{2}=\sqrt{\max{\text{eig}(\mathbf{A}^{T}\mathbf{A})}}$$

* Además de estas, existe la norma de Frobenius:

$$
\left\Vert \mathbf{A} \right\Vert_{F} = \sqrt{\sum_{i = 1}^{m} \sum_{j = 1}^{n} |a_{ij}^2|} = \sqrt{\text{tr}\,\left(\mathbf{A}^{*}\mathbf{A}\right)} = \sqrt{\sum_{i = 1}^{\min{m, n}} \sigma_{i}^2}
$$

donde $\mathbf{A^{*}}$ es la traspuesta conjugada y $\sigma_{i}$ son los valores singulares de $\mathbf{A}$

## Proyecto de tema

Hasta ahora, se ha trabajado de manera sencilla y práctica el manejo de la norma para las matrices. A manera de cimentar tus conocimientos y desarrollar tu pensamiento computacional, la tarea a elaborar esta ocasión es la siguiente:

Escribir la función ```normMat```, la cual recibe como parámetro una matriz de $n \times n$. Dicha función debe calcular la norma-2 para matrices siguiendo el ejemplo 2 del notebook de clase o siguiendo alguno de los métodos que se presentan en dicho notebook. Recuerda evaluar las condiciones necesarias para llevar a cabo este cálculo.

**NOTA**: Está permitido el uso del módulo ```numpy``` para el la manipulación de arreglo/matrices. No está permitido el uso de las funciones ```np.linalg.norm()``` y ```np.linalg.eigvals```. El cálculo de los eigenvalores debe ser explícito en la función.

<!-- grade_id: cell-51dcbfcf368e0cfd, points: 0, grade: false, solution: true -->
```python
import numpy as np

def normMat(matriz):
    
    #BEGIN SOLUTION
    n, m = matriz.shape
    if n == m:
        norma = np.sqrt(np.sum(matriz**2))
    else:
        return 0
    #END SOLUTION
    
    return norma
```

<!-- grade_id: cell-c21d43e621e9f322, points: 100, grade: true, solution: false -->
```python
"""Prueba: Evaluación de la función normMat"""
from numpy.testing import assert_approx_equal
from numpy.linalg import norm

# Norma-2 de matrices
# Casos válidos
assert_approx_equal(normMat(np.array([[2, 5], [1, 3]])), norm(np.array([[2, 5], [1, 3]])), 2)
print("Norma-2 de la matriz\n", np.array([[2, 5], [1, 3]]), "\nes:", normMat(np.array([[2, 5], [1, 3]])), "\n")

assert_approx_equal(normMat(np.array([[3, 7, 10], [20, 40, 10], [4, 5, 2]])), norm(np.array([[3, 7, 10], [20, 40, 10], [4, 5, 2]])), 2)
print("Norma-2 de la matriz\n", np.array([[3, 7, 10], [20, 40, 10], [4, 5, 2]]), "\nes:", 
      normMat(np.array([[3, 7, 10], [20, 40, 10], [4, 5, 2]])), "\n")

assert_approx_equal(normMat(np.array([[0, 1, 2, -1], [1, -1, -3, 1], [4, 0, 0, 1], [2, 3, 8, -2]])), 
                    norm(np.array([[0, 1, 2, -1], [1, -1, -3, 1], [4, 0, 0, 1], [2, 3, 8, -2]])), 2)
print("Norma-2 de la matriz\n", np.array([[0, 1, 2, -1], [1, -1, -3, 1], [4, 0, 0, 1], [2, 3, 8, -2]]), "\nes:", 
      normMat(np.array([[0, 1, 2, -1], [1, -1, -3, 1], [4, 0, 0, 1], [2, 3, 8, -2]])), "\n")
```